Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Mohlatlego Nakeng"
StudentNumber = "13413351"
```

---

## COS 802 Homework 1 - Part 1
## NLP Basics - Getting features


**You will learn how to:**
- Use Bag of Words sklearn NLP tools specifically Term Frequency


## 1 Packages ##

Scikit-Learn for text Analysis
- [sklearn](https://scikit-learn.org/stable/tutorial/text_analytics/working_with_text_data.html): Scikit-Learn Working With Text Data
- [matplotlib](http://matplotlib.org) is a library for plotting graphs in Python.

```python
import pandas as pd
import numpy as np
from sklearn.feature_extraction.text import CountVectorizer, TfidfVectorizer # Bag of Words: Term Frequencey and TFIDF
```

## 2 Let's get some data

The data we will use for this homework is the [20 Newsgroups Dataset](http://qwone.com/~jason/20Newsgroups/).

> The 20 Newsgroups data set is a collection of approximately 20,000 newsgroup documents, partitioned (nearly) evenly across 20 different newsgroups. To the best of my knowledge, it was originally collected by Ken Lang, probably for his Newsweeder: Learning to filter netnews paper, though he does not explicitly mention this collection. The 20 newsgroups collection has become a popular data set for experiments in text applications of machine learning techniques, such as text classification and text clustering.


```python
from sklearn.datasets import fetch_20newsgroups
newsgroups_train = fetch_20newsgroups(subset='train')
```

Lets look at a bit of this data. First lets put the data in a Pandas Data Frame.

```python
df_newsgroups_train = pd.DataFrame({'text':newsgroups_train['data'],
                                   'target': newsgroups_train['target']})
```

```python
df_newsgroups_train.head()
```

Output:
```
text  target
0  From: lerxst@wam.umd.edu (where's my thing)\nS...       7
1  From: guykuo@carson.u.washington.edu (Guy Kuo)...       4
2  From: twillis@ec.ecn.purdue.edu (Thomas E Will...       4
3  From: jgreen@amber (Joe Green)\nSubject: Re: W...       1
4  From: jcm@head-cfa.harvard.edu (Jonathan McDow...      14
```

```python
# Some information about the data
df_newsgroups_train.describe()
```

Output:
```
target
count  11314.000000
mean       9.293000
std        5.562719
min        0.000000
25%        5.000000
50%        9.000000
75%       14.000000
max       19.000000
```

```python
# Let us look at what the target names are
newsgroups_train_label_names = newsgroups_train['target_names']
newsgroups_train_label_names
```

Output:
```
['alt.atheism',
 'comp.graphics',
 'comp.os.ms-windows.misc',
 'comp.sys.ibm.pc.hardware',
 'comp.sys.mac.hardware',
 'comp.windows.x',
 'misc.forsale',
 'rec.autos',
 'rec.motorcycles',
 'rec.sport.baseball',
 'rec.sport.hockey',
 'sci.crypt',
 'sci.electronics',
 'sci.med',
 'sci.space',
 'soc.religion.christian',
 'talk.politics.guns',
 'talk.politics.mideast',
 'talk.politics.misc',
 'talk.religion.misc']
```

## Q1: Use Count Vectorizer with unigrams

**Hint**: Use the `CountVectorizer` with default settings. You can also check the tutorial on the Scikit-learn website that I linked to earlier.

<!-- grade_id: count_vectorizer, points: 0, grade: false, solution: true -->
```python
def initialise_term_frequency_vectorizer(data):
    # Must return X- transformed text data, vectorizer - vecrtorizer that is fit on the data
    # YOUR CODE HERE
    vectorizer_tf = CountVectorizer() #Initializing the CountVectorizer
    X = vectorizer_tf.fit_transform(data) #fitting the N-grams
  
    return X, vectorizer_tf
```

<!-- grade_id: count_vectorizer_test, points: 6, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
X, vectorizer_tf = initialise_term_frequency_vectorizer(df_newsgroups_train.text)
assert len(vectorizer_tf.vocabulary_) == 130107
assert len(vectorizer_tf.vocabulary_) == X.shape[1]
```

```python
X
```

Output:
```
<11314x130107 sparse matrix of type '<class 'numpy.int64'>'
	with 1787565 stored elements in Compressed Sparse Row format>
```

## Q2: Finding the top words [Uni-Grams]

In this section, you will
* Create a helper function to go from index in the vocabulary to a word
* Count the number of occurences of each word [identify the most used word]

Remember that we are using term frequency for representing our word tokens. It means that every row in **X**, contains the counts of the specific word in that document. Every row is a document. Every column is a word. You can find the total number of uses of a word by summing by row.

<!-- grade_id: cell-30ec7be7b35d3307, points: 0, grade: false, solution: true -->
```python
def get_id2word(vocabulary):
    id2word = {}
    for key in vocabulary.keys():
        # YOUR CODE HERE
        id2word = key
                      
    return id2word
```

```python
id2word
```

Output:
```
'willow'
```

```python
int[vectorizer_tf.vocabulary_[id2word]]
```

Output:
```
[124370]
```

<!-- grade_id: cell-926d40c25a006eb9, points: 5, grade: true, solution: false -->
```python
id2word = get_id2word(vectorizer_tf.vocabulary_)

assert type(id2word[0]) == str 
hello_index = vectorizer_tf.vocabulary_['hello']
assert id2word[hello_index] == 'hello'

```

Output:
```
IndexError                                Traceback (most recent call last)
<ipython-input-204-fb4f31d3a31d> in <module>
      3 assert type(id2word[0]) == str
      4 hello_index = vectorizer_tf.vocabulary_['hello']
----> 5 assert id2word[hello_index] == 'hello'

IndexError: string index out of range
```

<!-- grade_id: cell-e6e852225be96389, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2d8935070c6722c4, points: 3, grade: true, solution: false -->
```python
assert token_counts.shape[1] == X.shape[1]
assert token_counts.shape[0] == 1
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-100-ba74a1b30bde> in <module>
----> 1 assert token_counts.shape[1] == X.shape[1]
      2 assert token_counts.shape[0] == 1

NameError: name 'token_counts' is not defined
```

<!-- grade_id: cell-50c7d795206f415d, points: 3, grade: true, solution: false -->
```python
list_token_counts = token_counts.tolist()[0]
assert list_token_counts[0] ==1534
assert list_token_counts[1] == 953
assert list_token_counts[10] == 2
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-101-4c1fd5b04e3f> in <module>
----> 1 list_token_counts = token_counts.tolist()[0]
      2 assert list_token_counts[0] ==1534
      3 assert list_token_counts[1] == 953
      4 assert list_token_counts[10] == 2

NameError: name 'token_counts' is not defined
```

```python
sorted_index = np.argsort(list_token_counts)[::-1]
print("Highest occuring word:%s - count:%s" % (id2word[sorted_index[0]], list_token_counts[sorted_index[0]]))
print("2nd highest occuring word:%s - count:%s" % (id2word[sorted_index[1]], list_token_counts[sorted_index[1]]))
print("3rd highest occuring word:%s - count:%s" % (id2word[sorted_index[2]], list_token_counts[sorted_index[2]]))

print("Lowest occuring word:%s - count:%s" % (id2word[sorted_index[-1]], list_token_counts[sorted_index[-1]]))
print("2nd lowest occuring word:%s - count:%s" % (id2word[sorted_index[-2]], list_token_counts[sorted_index[-2]]))
print("3rd lowest occuring word:%s - count:%s" % (id2word[sorted_index[-3]], list_token_counts[sorted_index[-3]]))
```

What do you note about the top words?

<!-- grade_id: stop-words, points: 7, grade: true, solution: true -->
YOUR ANSWER HERE

## Q3: Finding the top words [n-grams]

In this section, you will be repeating Q2 tasks, but now trying to use bigrams. Look at the documentation for the ngram_range variable for CountVectorizer

<!-- grade_id: cell-4fb7ef339af2691b, points: 0, grade: false, solution: true -->
```python
def initialise_term_frequency_vectorizer_ngram(data, ngram = 1):
    # Must return X- transformed text data, vectorizer - vectorizer that is fit on the data, must return for the n_gram specified
    # YOUR CODE HERE
    raise NotImplementedError()
    return X, vectorizer_tf
```

<!-- grade_id: cell-7259a3c286adbd0d, points: 3, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
X_bigram, vectorizer_tf = initialise_term_frequency_vectorizer_ngram(df_newsgroups_train.text, ngram=2)
assert len(vectorizer_tf.vocabulary_) == 1051696
assert len(vectorizer_tf.vocabulary_) == X_bigram.shape[1]
```

```python
bigram_counts = np.sum(X_trigram,axis=0).tolist()[0]
id2word = get_id2word(vectorizer_tf.vocabulary_)
```

```python
sorted_index = np.argsort(bigram_counts)[::-1]
print("Highest occuring word:%s - count:%s" % (id2word[sorted_index[0]], bigram_counts[sorted_index[0]]))
print("2nd highest occuring word:%s - count:%s" % (id2word[sorted_index[1]], bigram_counts[sorted_index[1]]))
print("3rd highest occuring word:%s - count:%s" % (id2word[sorted_index[2]], bigram_counts[sorted_index[2]]))

```

What do you note about the top n-grams

<!-- grade_id: cell-71c70c7458b660e3, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE